---
title: Template Chapter 1
description: This is a template chapter.
---

## Ex 1.1

```yaml
type: NormalExercise
lang: sql
xp: 100
skills: 1
key: 8e774c2a97
```

Do some data science.

`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```

`@sample_code`

```{sql}

```

`@solution`

```{python}

```

`@sct`

```{sql}

```
